
# Assignment No 3b
###### *Course Instructors: Dr. Hassan Raza, Mr. Ahmad Raza*
----
## Goal

Your goal in this part of assigment is to implement a Decision Tree Classifier for categorical variables.

**Note** Please note that you are allowed to use only those libraries which we have discussed in the class, i.e. numpy, scipy, pandas.

## Submission Instructions
You are required to submit the original notebook file on the google classroom (with .ipynb extension), with complete set of outputs. Students failing to do so will get zero marks. 

*Please read each step carefully and understand it fully before proceeding with code writing*

## Plagiarism
Any form of plagiarism will not be tolerated and result in 0 marks.



### Decision Tree Classifier

Now in this assignment we will be implementing the Decision Classifier for both Continuous and Categorical attributes.

Decision tree can be built by using any of the following split criterias, namely:
 - Information Gain
 - Gini Index
 - CART 

However, you are required here to implement the decision tree with information gain as splitting criterion.

Remember in my code i am not looking for maximizing the information gain, instead i am looking for minimizing the split entropy. Recall,
$$Information Gain  = H(D) - H(D_Y,D_N)$$

Where,

$H(D)$ is the data set entroy and $H(D_Y,D_N)$ is split entropy. Since $H(D)$ is constant for the given dataset so maximizing the entropy is equal to minimizing the split entropy and that is what is being represented in my code outputs.

```python
%pylab inline
import scipy.stats
from collections import defaultdict  # default dictionary 
```

Output:
```
%pylab is deprecated, use %matplotlib inline and import the required libraries.
Populating the interactive namespace from numpy and matplotlib
```

```python
#Cutomize the Matplotlib for beautiful plots...
# import dmStyle
# dmStyle.customize_mpl()

```

```python
import pandas as pd
import tools as t # set of tools for plotting, data splitting, etc..
import numpy as np
```

```python
def getSplits(categories):
    '''
        function returns list of splits for the given list of categorical variables...
        
        Input:
        ------------
            categories: a list of unique categories...
        
        Return:
        ------------
            list of splits(tuples) for given list of categorical variables. Each pair of sublists
            defines the left and right splits, e.g. This list
            [('y', 'f'), ('s', 'g'), ('f'), ('y', 's', 'g')]
            
            defines two splits with each pair representing a different split.
        Examples:
        ------------
        splits=getSplits(['a1','a2','a3','a4']) will return 
        [('a1', 'a2', 'a4'), ('a3',), ('a2', 'a4'), 
        ('a1', 'a3'), ('a1', 'a4'), ('a3', 'a2'), ('a1', 'a3', 'a2'), 
        ('a4',), ('a3', 'a4'), ('a1', 'a2'), ('a1',), ('a3', 'a2', 'a4'), ('a2',), ('a1', 'a3', 'a4')]

            
    '''
    categories=set(categories)
    tsplits=t.get_powerset(categories,len(categories)-1) # get all the power sets with the given cardinality...
    flist=[]
    for s in tsplits:

        if not s in flist:
            r=categories.difference(s)
    #         print s, tuple(r)
            flist.append(s)
            flist.append(r)

    olist=[]

    for s in flist:
        ilist=[]

        for k in s:
            ilist.append(k)
        olist.append(tuple(ilist))

#     print olist
    
    return olist
```

```python
splits=getSplits(['a1','a2','a3','a4'])
print (splits,len(splits))
```

Output:
```
[('a2', 'a4'), ('a1', 'a3'), ('a1', 'a4'), ('a2', 'a3'), ('a4',), ('a2', 'a1', 'a3'), ('a3',), ('a1', 'a4', 'a2'), ('a1',), ('a2', 'a4', 'a3'), ('a2',), ('a1', 'a4', 'a3'), ('a1', 'a2'), ('a4', 'a3')] 14
```

<!-- grade_id: node, points: 0, grade: false, solution: true -->
```python
class Node:
    def __init__(self,purity,klasslabel='',score=0,split=[],fidx=-1):
        '''
            purity: purity level at which to stop
            klasslabel: klasslabel of the node, (for leaf node)
            score: information gain of the newly added node
            split: splitting threshold
            fidx: feature index            
        '''
        self.lchild=None
        self.rchild=None
        self.klasslabel=klasslabel        
        self.split=split
        self.score=score
        self.fidx=fidx
        self.purity=purity
        self.ftype = 'categorical' if isinstance(self.split, (tuple, str, np.str_)) else 'continuous'


        
    def set_childs(self,lchild,rchild):
        # YOUR CODE HERE
        #raise NotImplementedError()
        self.lchild = lchild
        self.rchild = rchild
    def isleaf(self):
        # YOUR CODE HERE
        #raise NotImplementedError()
        if(self.lchild is None and self.rchild is None):
            return True
    def isless_than_eq(self, X):
        # YOUR CODE HERE
        #raise NotImplementedError()
        return X[self.fidx] <= self.split
    def get_str(self):        
        if self.isleaf():
            return 'C(class={},Purity={})'.format(self.klasslabel,self.purity)
        else:
            return 'I(Fidx={},Score={},Split={})'.format(self.fidx,self.score,self.split)
    
```

```python

```

```python
# A placeholder class 
# TODO: You have to implement the following class, remember from the lectures that you will 
# need to build a model for each different class you are trying to identify...
```

<!-- grade_id: tree, points: 0, grade: false, solution: true -->
```python
### import pdb
## Your code goes here...
# You might need to define auxliary classes for composition.. ?
class DecisionTree:
    ''' Implements the Decision Tree For Classification... '''
    def __init__(self, purityp, exthreshold, maxdepth=10):        
        self.purity = purityp
        self.exthreshold = exthreshold
        self.maxdepth = maxdepth
        self.tree = None
        
        pass
    def train(self, X, Y):
        ''' Train Decision Tree using the given 
            X [m x d] data matrix and Y labels matrix
            
            Input:
            ------
            X: [m x d] a data matrix of m d-dimensional examples.
            Y: [m x 1] a label vector.
            
            Returns:
            -----------
            Nothing
            '''
        nexamples, nfeatures = X.shape
        # # now go and train a model for each class...
        # YOUR CODE HERE
        #raise NotImplementedError()
        self.tree = self.build_tree(X, Y, depth=0)
        
    def build_tree(self, X, Y, depth):
        """ 
            Function is used to recursively build the decision Tree 
          
            Input
            -----
            X: [m x d] a data matrix of m d-dimensional examples.
            Y: [m x 1] a label vector.
            
            Returns
            -------
            root node of the built tree...
        """
       # print(f"X.shape: {X.shape}")  # Debugging line
        
        nexamples, nfeatures = X.shape
        
        
        # YOUR CODE HERE
        #raise NotImplementedError()
        
        if depth >= self.maxdepth or nexamples<=self.exthreshold or len(set(Y)) == 1:
            klasslabel = max(set(Y) , key=list(Y).count)
            return Node(purity=self.purity , klasslabel=klasslabel )
        

        best_score , best_split , best_lidx , best_ridx = float('-inf') , None , None , None
        best_feature = None
        
        for i in range(nfeatures):
            feat = X[: , i]
            split, entropy, lidx, ridx , gain = self.evaluate_categorical_attribute(feat, Y)
            
            if len(lidx) == 0 or len(ridx) == 0:
                continue  # Skip this split and try another one
            
            
            if gain > best_score:
                best_score = gain
                best_split = split
                best_feature = i
                best_lidx = lidx
                best_ridx = ridx
                
        if best_lidx is None or best_ridx is None:
            klasslabel = max(set(Y), key=list(Y).count)  # Majority class label
            return Node(purity=self.purity, klasslabel=klasslabel)
        
        node = Node(purity=self.purity, score=best_score, split=best_split, fidx=best_feature)
        

        # Recursively build the tree by splitting into child nodes
        node.set_childs(
            self.build_tree(X[best_lidx], Y[best_lidx], depth + 1),
            self.build_tree(X[best_ridx], Y[best_ridx], depth + 1)
        )
        
        return node
        
        
    def test(self, X):
        
        ''' Test the trained classifiers on the given set of examples 
        
                   
            Input:
            ------
            X: [m x d] a data matrix of m d-dimensional test examples.
           
            Returns:
            -----------
                pclass: the predicted class for each example, i.e. to which it belongs
        '''
        
        nexamples, nfeatures = X.shape
        pclasses = self.predict(X)
        
        # your code go here...
        
        return np.array(pclasses)
    
    def information_gain(self, parent, left_values ,  right_values):    
        prob_left = len(left_values) / len(parent)
        prob_right = len(right_values) / len(parent)
        parent_impurity = self.impurity(parent)
        left_impurity = self.impurity(left_values)
        right_impurity = self.impurity(right_values)
        
        split_entropy =  prob_left*left_impurity+prob_right*right_impurity
        value = parent_impurity - split_entropy
        
        
        return value , split_entropy
        # info_gain = parent_impurity - (prob_left * left_impurity + prob_right * right_impurity)
        # return info_gain
        
        
    def impurity(self , y):
        values , counts = np.unique(y , return_counts=True)
        prob = counts / len(y)
        return -np.sum(prob * np.log2(prob))
    
    
    def evaluate_categorical_attribute(self, feat, Y):
            """
            Evaluates the cateogrical attribute for all possible split points, i.e. 2^(m-1)-2 for
            possible feature selection

            Input:
            ---------
            feat: a categorical feature
            Y: labels

            Returns:
            ----------
            v: splitting threshold
            score: splitting score
            Xlidx: Index of examples belonging to left child node
            Xridx: Index of examples belonging to right child node

            """
            categories = set(feat)
            splits = getSplits(categories) if len(categories) > 1 else tuple(categories)
           # freq = scipy.stats.itemfreq(Y)
            best_gain = float('-inf')
            min_entropy = 0
            best_split, best_lidx, best_ridx = None, None, None
            for split in splits:
                # for index , value in enumerate(feat):
                #     if value in i:
                #         lidx.append(index)
                #     else:
                #         ridx.append(index)
                Xlidx = [i for i, val in enumerate(feat) if val in split]
                Xridx = [i for i, val in enumerate(feat) if val not in split]     
                
                left_values = Y[Xlidx]
                right_values = Y[Xridx]   
                info_gain , split_entropy = self.information_gain(Y ,  left_values  , right_values)
                
                if info_gain > best_gain:
                    best_gain = info_gain
                    best_split = split
                    min_entropy = split_entropy
                    best_lidx = Xlidx
                    best_ridx = Xridx
        
            # YOUR CODE HERE
            #raise NotImplementedError()
            return best_split , min_entropy , best_lidx , best_ridx , best_gain
            #return csplit, mingain, Xlidx, Xridx

        
    def evaluate_numerical_attribute(self, feat, Y):
        '''
            Evaluates the numerical attribute for all possible split points for
            possible feature selection
            
            Input:
            ---------
            feat: a contiuous feature
            Y: labels
            
            Returns:
            ----------
            v: splitting threshold
            score: splitting score
            Xlidx: Index of examples belonging to left child node
            Xridx: Index of examples belonging to right child node
            
        '''
        
        classes = np.unique(Y)
        nclasses = len(self.classes)
        f = np.sort(feat)
        
        # YOUR CODE HERE
        raise NotImplementedError()
        
        return split, mingain, Xlidx, Xridx
    def predict(self, X):
        
        """
        Test the trained classifiers on the given example X
        
                   
            Input:
            ------
            X: [1 x d] a d-dimensional test example.
           
            Returns:
            -----------
                pclass: the predicted class for the given example, i.e. to which it belongs
        """
        z = []
        
        for idx in range(X.shape[0]):
            
            z.append(self._predict(self.tree, X[idx, :]))
        
        return z 

    
    
    def _predict(self, node, X):
        
        if node.isleaf():
            return node.klasslabel
        if X[node.fidx] in node.split:
            return self._predict(node.lchild, X)
        else:
            return self._predict(node.rchild, X)
        
        # YOUR CODE HERE
        #raise NotImplementedError()

    def __str__(self):
        str = '---------------------------------------------------'
        str += '\n A Decision Tree With Depth={}'.format(self.find_depth())
        str += self.__print(self.tree)
        str += '\n---------------------------------------------------'
        return str  # self.__print(self.tree)        
        
     
   
    def find_depth(self):
        return self._find_depth(self.tree)
    def _find_depth(self, node):
        if not node:
            return
        if node.isleaf():
            return 1
        else:
            return max(self._find_depth(node.lchild), self._find_depth(node.rchild)) + 1
    def __print(self, node, depth=0):
        
        ret = ""

        # Print right branch
        if node.rchild:
            ret += self.__print(node.rchild, depth + 1)

        # Print own value
        
        ret += "\n" + ("    "*depth) + node.get_str()

        # Print left branch
        if node.lchild:
            ret += self.__print(node.lchild, depth + 1)
        
        return ret
```

```python
# A placeholder class 
# TODO: You have to implement the following class, remember from the lectures that you will 
# need to build a model for each different class you are trying to identify...
```

```python

```

# Lets test our code for the given example in the book.

```python
#load the Iris dataset
tdata=pd.read_csv('./iris.data')
tdata.columns=['SepalLength','SepalWidth','PetalLength','PetalWidth','Class']
```

```python
tx=tdata['SepalLength'].dropna()
tx[(tdata['SepalLength']>=4.3) & (tdata['SepalLength']<=5.2)]='a1'
tx[(tdata['SepalLength']>5.2) & (tdata['SepalLength']<=6.1)]='a2'
tx[(tdata['SepalLength']>6.1) & (tdata['SepalLength']<=7.0)]='a3'
tx[(tdata['SepalLength']>7.0) & (tdata['SepalLength']<=7.9)]='a4'
```

Output:
```
C:\Users\muham\AppData\Local\Temp\ipykernel_17524\1702864819.py:2: FutureWarning: Setting an item of incompatible dtype is deprecated and will raise an error in a future version of pandas. Value 'a1' has dtype incompatible with float64, please explicitly cast to a compatible dtype first.
  tx[(tdata['SepalLength']>=4.3) & (tdata['SepalLength']<=5.2)]='a1'
```

```python
print (tx.values)
Y=tdata['Class'].dropna()
Y[Y=='Iris-virginica']='Iris-versicolor'
Y=Y.values
print (Y)
```

Output:
```
['a1' 'a1' 'a1' 'a1' 'a2' 'a1' 'a1' 'a1' 'a1' 'a2' 'a1' 'a1' 'a1' 'a2'
 'a2' 'a2' 'a1' 'a2' 'a1' 'a2' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1'
 'a1' 'a1' 'a2' 'a1' 'a2' 'a1' 'a1' 'a2' 'a1' 'a1' 'a1' 'a1' 'a1' 'a1'
 'a1' 'a1' 'a1' 'a1' 'a1' 'a2' 'a1' 'a3' 'a3' 'a3' 'a2' 'a3' 'a2' 'a3'
 'a1' 'a3' 'a1' 'a1' 'a2' 'a2' 'a2' 'a2' 'a3' 'a2' 'a2' 'a3' 'a2' 'a2'
 'a2' 'a3' 'a2' 'a3' 'a3' 'a3' 'a3' 'a2' 'a2' 'a2' 'a2' 'a2' 'a2' 'a2'
 'a2' 'a3' 'a3' 'a2' 'a2' 'a2' 'a2' 'a2' 'a1' 'a2' 'a2' 'a2' 'a3' 'a1'
 'a2' 'a3' 'a2' 'a4' 'a3' 'a3' 'a4' 'a1' 'a4' 'a3' 'a4' 'a3' 'a3' 'a3'
 'a2' 'a2' 'a3' 'a3' 'a4' 'a4' 'a2' 'a3' 'a2' 'a4' 'a3' 'a3' 'a4' 'a3'
 'a2' 'a3' 'a4' 'a4' 'a4' 'a3' 'a3' 'a2' 'a4' 'a3' 'a3' 'a2' 'a3' 'a3'
 'a3' 'a2' 'a3' 'a3' 'a3' 'a3' 'a3' 'a3' 'a2']
['Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa'
 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa'
 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa' 'Iris-setosa'
 'Iris-setosa' 'Iris-setosa' '
```

```python
%pdb off
dt=DecisionTree(0.95,5,5)
split, gain, Xlidx, Xridx , _  = dt.evaluate_categorical_attribute(tx,Y)
print ()
print (split, gain)
```

Output:
```
Automatic pdb calling has been turned OFF

('a1',) 0.5107023503160576
```

<!-- grade_id: split_gain, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_almost_equal, assert_almost_equals

%pdb off
dt=DecisionTree(0.95,5,5)
split, gain, Xlidx, Xridx , _ = dt.evaluate_categorical_attribute(tx,Y)


assert split[0] == 'a1', f"Expected 'a1', but got {split[0]}"
assert abs(gain - 0.51) < 0.1, f"Expected gain approximately 0.51, but got {gain}"
```

Output:
```
Automatic pdb calling has been turned OFF
```

# Now lets test on [Mushroom](https://archive.ics.uci.edu/ml/machine-learning-databases/mushroom/agaricus-lepiota.names) dataset


This data set includes descriptions of hypothetical samples corresponding to 23 species of gilled mushrooms in the Agaricus and Lepiota Family (pp. 500-525).  Each species is identified as definitely edible, definitely poisonous, or of unknown edibility and not recommended.  This latter class was combined with the poisonous one.  The Guide clearly states that there is no simple rule for determining the edibility of a mushroom; no rule like ``leaflets three, let it be'' for Poisonous Oak and Ivy.


- Number of Instances: 8124

- Number of Attributes: 22 (all nominally valued)

- Attribute Information: (classes: edible=e, poisonous=p)
         1. cap-shape:                bell=b,conical=c,convex=x,flat=f, knobbed=k,sunken=s
         2. cap-surface:              fibrous=f,grooves=g,scaly=y,smooth=s
         3. cap-color:                brown=n,buff=b,cinnamon=c,gray=g,green=r, pink=p,purple=u,red=e,white=w,yellow=y
         4. bruises?:                 bruises=t,no=f
         5. odor:                     almond=a,anise=l,creosote=c,fishy=y,foul=f, musty=m,none=n,pungent=p,spicy=s
         6. gill-attachment:          attached=a,descending=d,free=f,notched=n
         7. gill-spacing:             close=c,crowded=w,distant=d
         8. gill-size:                broad=b,narrow=n
         9. gill-color:               black=k,brown=n,buff=b,chocolate=h,gray=g, green=r,orange=o,pink=p,purple=u,red=e,white=w,yellow=y
        10. stalk-shape:              enlarging=e,tapering=t
        11. stalk-root:               bulbous=b,club=c,cup=u,equal=e, rhizomorphs=z,rooted=r,missing=?
        12. stalk-surface-above-ring: fibrous=f,scaly=y,silky=k,smooth=s
        13. stalk-surface-below-ring: fibrous=f,scaly=y,silky=k,smooth=s
        14. stalk-color-above-ring:   brown=n,buff=b,cinnamon=c,gray=g,orange=o, pink=p,red=e,white=w,yellow=y
        15. stalk-color-below-ring:   brown=n,buff=b,cinnamon=c,gray=g,orange=o, pink=p,red=e,white=w,yellow=y
        16. veil-type:                partial=p,universal=u
        17. veil-color:               brown=n,orange=o,white=w,yellow=y
        18. ring-number:              none=n,one=o,two=t
        19. ring-type:                cobwebby=c,evanescent=e,flaring=f,large=l, none=n,pendant=p,sheathing=s,zone=z
        20. spore-print-color:        black=k,brown=n,buff=b,chocolate=h,green=r, orange=o,purple=u,white=w,yellow=y
        21. population:               abundant=a,clustered=c,numerous=n, scattered=s,several=v,solitary=y
        22. habitat:                  grasses=g,leaves=l,meadows=m,paths=p, urban=u,waste=w,woods=d

- Missing Attribute Values: 2480 of them (denoted by "?"), all for  attribute #11.

- Class Distribution: 
        --    edible: 4208 (51.8%)
        -- poisonous: 3916 (48.2%)
        --     total: 8124 instances

```python
#load the data set
data=pd.read_csv('./mushrooms.csv')
data.columns=columns=['class', 'cap-shape','cap-surface','cap-color','bruises?','odor','gill-attachment','gill-spacing','gill-size','gill-color','stalk-shape','stalk-root','stalk-surface-above-ring','stalk-surface-below-ring','stalk-color-above-ring','stalk-color-below-ring','veil-type','veil-color','ring-number','ring-type','spore-print-color','population','habitat']
print (data.describe())
```

Output:
```
class cap-shape cap-surface cap-color bruises?  odor gill-attachment  \
count   8124      8124        8124      8124     8124  8124            8124   
unique     2         6           4        10        2     9               2   
top        e         x           y         n        f     n               f   
freq    4208      3656        3244      2284     4748  3528            7914   

       gill-spacing gill-size gill-color  ... stalk-surface-below-ring  \
count          8124      8124       8124  ...                     8124   
unique            2         2         12  ...                        4   
top               c         b          b  ...                        s   
freq           6812      5612       1728  ...                     4936   

       stalk-color-above-ring stalk-color-below-ring veil-type veil-color  \
count                    8124                   8124      8124       8124   
unique                      9                      9         1          4   
top      
```

```python
data.tail()
```

Output:
```
class cap-shape cap-surface cap-color bruises? odor gill-attachment  \
8119     e         k           s         n        f    n               a   
8120     e         x           s         n        f    n               a   
8121     e         f           s         n        f    n               a   
8122     p         k           y         n        f    y               f   
8123     e         x           s         n        f    n               a   

     gill-spacing gill-size gill-color  ... stalk-surface-below-ring  \
8119            c         b          y  ...                        s   
8120            c         b          y  ...                        s   
8121            c         b          n  ...                        s   
8122            c         n          b  ...                        k   
8123            c         b          y  ...                        s   

     stalk-color-above-ring stalk-color-below-ring veil-type veil-color  \
8119                      o             
```

```python
# Get your data in matrix
X=np.asarray(data.iloc[:,1:].dropna())
Y=np.asarray(data.iloc[:,0].dropna())
print (" Data Set Dimensions=", X.shape, " True Class labels dimensions", Y.shape )  

```

Output:
```
Data Set Dimensions= (8124, 22)  True Class labels dimensions (8124,)
```

```python
# Split your data into training and test-set... 
# see the documentation of split_data in tools for further information...
Xtrain,Ytrain,Xtest,Ytest=t.split_data(X,Y)

print (" Training Data Set Dimensions=", Xtrain.shape, "Training True Class labels dimensions", Ytrain.shape )  
print (" Test Data Set Dimensions=", Xtest.shape, "Test True Class labels dimensions", Ytrain.shape  ) 

```

Output:
```
Training Data Set Dimensions= (5687, 22) Training True Class labels dimensions (5687,)
 Test Data Set Dimensions= (2437, 22) Test True Class labels dimensions (5687,)
```

```python
# Lets train a Decision Tree Classifier two features
# %pdb
feat=np.arange(2)
dt=DecisionTree(0.95,5,10)
dt.train(Xtrain[:,feat],Ytrain)
```

Output:
```
Cell In[8], line 85, in DecisionTree.build_tree(self, X, Y, depth)
     80 node = Node(purity=self.purity, score=best_score, split=best_split, fidx=best_feature)
     83 # Recursively build the tree by splitting into child nodes
     84 node.set_childs(
---> 85     self.build_tree(X[best_lidx], Y[best_lidx], depth + 1),
     86     self.build_tree(X[best_ridx], Y[best_ridx], depth + 1)
     87 )
     89 return node

Cell In[8], line 77, in DecisionTree.build_tree(self, X, Y, depth)
     74         best_ridx = ridx
     76 if best_lidx is None or best_ridx is None:
---> 77     klasslabel = max(set(Y), key=list(Y).count)  # Majority class label
     78     return Node(purity=self.purity, klasslabel=klasslabel)
     80 node = Node(purity=self.purity, score=best_score, split=best_split, fidx=best_feature)

TypeError: max() got an unexpected keyword argument 'key'
```

```python
print (dt)
```

Output:
```
---------------------------------------------------
 A Decision Tree With Depth=7
                C(class=e,Purity=0.95)
            I(Fidx=1,Score=0.0011053721757735002,Split=('s',))
                C(class=e,Purity=0.95)
        I(Fidx=0,Score=0.010159885040279482,Split=('s',))
            C(class=e,Purity=0.95)
    I(Fidx=1,Score=0.013760273096525943,Split=('y',))
        C(class=e,Purity=0.95)
I(Fidx=0,Score=0.03096237523730494,Split=('k', 'f', 'x', 'c'))
            C(class=e,Purity=0.95)
        I(Fidx=0,Score=0.005869587232921858,Split=('f', 'x'))
                C(class=e,Purity=0.95)
            I(Fidx=0,Score=0.001623251972961759,Split=('x',))
                C(class=e,Purity=0.95)
    I(Fidx=1,Score=0.031112422124874972,Split=('y', 'g', 's'))
                C(class=p,Purity=0.95)
            I(Fidx=1,Score=0.04445650091659803,Split=('y', 'g'))
                    C(class=p,Purity=0.95)
                I(Fidx=0,Score=0.001896107945202341,Split=('k',))
                       
```

```python
pclasses=dt.predict(Xtest[:,feat])
#Lets see how good we are doing, by finding the accuracy on the test set..
print (np.sum(pclasses==Ytest))
print ("Accuracy = ", np.sum(pclasses==Ytest)/float(Ytest.shape[0]))
```

Output:
```
1509
Accuracy =  0.6192039392695937
```

<!-- grade_id: acc, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_greater_equal
Xtrain,Ytrain,Xtest,Ytest=t.split_data(X,Y)
feat=np.arange(2)
dt=DecisionTree(0.95,5,10)
dt.train(Xtrain[:,feat],Ytrain)


pclasses=dt.predict(Xtest[:,feat])
acc = np.sum(pclasses==Ytest)/float(Ytest.shape[0])

assert_greater_equal(acc, 0.60)
```

Output:
```
X.shape: (5687, 2)
X.shape: (5354, 2)
X.shape: (3770, 2)
X.shape: (3224, 2)
X.shape: (1338, 2)
X.shape: (762, 2)
X.shape: (576, 2)
X.shape: (575, 2)
X.shape: (1, 2)
X.shape: (1886, 2)
X.shape: (981, 2)
X.shape: (905, 2)
X.shape: (546, 2)
X.shape: (302, 2)
X.shape: (244, 2)
X.shape: (241, 2)
X.shape: (3, 2)
X.shape: (1584, 2)
X.shape: (1537, 2)
X.shape: (824, 2)
X.shape: (713, 2)
X.shape: (47, 2)
X.shape: (333, 2)
X.shape: (99, 2)
X.shape: (234, 2)
X.shape: (21, 2)
X.shape: (213, 2)
X.shape: (173, 2)
X.shape: (40, 2)
```

<!-- grade_id: acc_all, points: 5, grade: true, solution: false -->
```python
from nose.tools import assert_greater_equal
Xtrain,Ytrain,Xtest,Ytest=t.split_data(X,Y)

dt=DecisionTree(0.95,5,10)
dt.train(Xtrain,Ytrain)

pclasses=dt.predict(Xtest)
acc = np.sum(pclasses==Ytest)/float(Ytest.shape[0])

assert_greater_equal(acc, 0.90)
```

Output:
```
X.shape: (5687, 22)
X.shape: (3033, 22)
X.shape: (52, 22)
X.shape: (2981, 22)
X.shape: (2910, 22)
X.shape: (2906, 22)
X.shape: (138, 22)
X.shape: (135, 22)
X.shape: (3, 22)
X.shape: (2768, 22)
X.shape: (4, 22)
X.shape: (71, 22)
X.shape: (43, 22)
X.shape: (28, 22)
X.shape: (2654, 22)
```

# Lets Train on all the features and for both the classes....

```python
# Split your data into training and test-set... 
# see the documentation of split_data in tools for further information...
Xtrain,Ytrain,Xtest,Ytest=t.split_data(X,Y)

print (" Training Data Set Dimensions=", Xtrain.shape, "Training True Class labels dimensions", Ytrain.shape )  
print (" Test Data Set Dimensions=", Xtest.shape, "Test True Class labels dimensions", Ytrain.shape   )

```

Output:
```
Training Data Set Dimensions= (5687, 22) Training True Class labels dimensions (5687,)
 Test Data Set Dimensions= (2437, 22) Test True Class labels dimensions (5687,)
```

```python
feat=np.arange(22)#[0, 1, 2, 3]
dt=DecisionTree(0.95,5,10)
dt.train(Xtrain[:,feat],Ytrain)
pclasses=dt.predict(Xtest[:,feat])
#Lets see how good we are doing, by finding the accuracy on the test set..
print (np.sum(pclasses==Ytest))
print ("Accuracy = ", np.sum(pclasses==Ytest)/float(Ytest.shape[0]))
```

Output:
```
X.shape: (5687, 22)
X.shape: (3035, 22)
X.shape: (51, 22)
X.shape: (2984, 22)
X.shape: (2907, 22)
X.shape: (2904, 22)
X.shape: (2903, 22)
X.shape: (133, 22)
X.shape: (132, 22)
X.shape: (1, 22)
X.shape: (2770, 22)
X.shape: (1, 22)
X.shape: (3, 22)
X.shape: (77, 22)
X.shape: (44, 22)
X.shape: (33, 22)
X.shape: (2652, 22)
2437
Accuracy =  1.0
```

```python
print (dt)
```

Output:
```
---------------------------------------------------
 A Decision Tree With Depth=8
    C(class=p,Purity=0.95)
I(Fidx=4,Score=0.8971345644457117,Split=('n', 'a', 'l'))
                C(class=p,Purity=0.95)
            I(Fidx=10,Score=0.9852281360342515,Split=('b',))
                C(class=e,Purity=0.95)
        I(Fidx=14,Score=0.05519404917848908,Split=('w', 'e', 'g', 'p', 'o'))
                C(class=p,Purity=0.95)
            I(Fidx=1,Score=0.010055781449585435,Split=('f', 'y', 's'))
                    C(class=p,Purity=0.95)
                I(Fidx=0,Score=0.003769377531839254,Split=('k', 'x', 'b', 's', 'f'))
                        C(class=e,Purity=0.95)
                    I(Fidx=21,Score=0.0015340114430465006,Split=('l',))
                            C(class=p,Purity=0.95)
                        I(Fidx=1,Score=0.06385368546122085,Split=('s',))
                            C(class=e,Purity=0.95)
    I(Fidx=19,Score=0.09422677099912931,Split=('r',))
        C(class=p,Purity=0.95)
-
```

What can you conclude ?
====================
Please write your observation....



# Cross-Validation

Until now we have been splitting the dataset into a training and test set rather randomly and were reporting a rather artifical performance. Now we are going to test our system exhaustively by making use of k-fold [cross validation](http://en.wikipedia.org/wiki/Cross-validation_%28statistics%29). 

Now go and tune your hyper-parameters (purity, exthreshold) to opitmize the performance.

```python
# Now lets cross validate for best paramters, and test the result...
# We will be training four different models on four different partitions of data set and 
# then will be reporting the mean accuracy of the four classifiers.

nfolds=4 # lets use four folds..
folds=t.generate_folds(X,Y,nfolds)
features=np.arange(22)#[0, 1, 2, 3] # features to use for our system
#now lets train and test on these folds...
totacc=[]
for k in range(nfolds):
    dt=DecisionTree(0.95,5,7)
    dt.train(folds[k][0][:,features],folds[k][1])
    pclasses=dt.predict(folds[k][2][:,features])
    acc=np.sum(pclasses==folds[k][3])/float(folds[k][3].shape[0])
    print (dt)
    print ("[Info] Fold {} Accuracy = {}".format(k+1, acc))
    totacc.append(acc)

print (totacc, '\n Mean Accuracy =', np.mean(totacc))
```

Output:
```
Generating CV data for 2 classes
X.shape: (6093, 22)
X.shape: (3251, 22)
X.shape: (59, 22)
X.shape: (3192, 22)
X.shape: (3112, 22)
X.shape: (3109, 22)
X.shape: (151, 22)
X.shape: (148, 22)
X.shape: (3, 22)
X.shape: (2958, 22)
X.shape: (3, 22)
X.shape: (80, 22)
X.shape: (50, 22)
X.shape: (30, 22)
X.shape: (2842, 22)
---------------------------------------------------
 A Decision Tree With Depth=7
    C(class=p,Purity=0.95)
I(Fidx=4,Score=0.8974380126101213,Split=('n', 'a', 'l'))
                C(class=p,Purity=0.95)
            I(Fidx=10,Score=0.954434002924965,Split=('b',))
                C(class=e,Purity=0.95)
        I(Fidx=14,Score=0.0455701849459859,Split=('w', 'e', 'g', 'p', 'o'))
                C(class=p,Purity=0.95)
            I(Fidx=1,Score=0.009120176679144509,Split=('f', 'y', 's'))
                    C(class=e,Purity=0.95)
                I(Fidx=21,Score=0.004224086026828071,Split=('l',))
                        C(class=p,Purity=0.95)
                    I(Fidx=1,Score=0
```

```python
# Now lets cross validate for best paramters, and test the result...
# We will be training four different models on four different partitions of data set and 
# then will be reporting the mean accuracy of the four classifiers.

nfolds=4 # lets use four folds..
folds=t.generate_folds(X,Y,nfolds)
features=[0,1, 2, 3] # features to use for our system
#now lets train and test on these folds...

#Lets perform the grid search...
purity=np.linspace(0.85,0.97,13)
nexamp=np.linspace(5,25,21)  

params=np.zeros((len(purity),len(nexamp)))
                   
for p in range(len(purity)):
    for n in range(len(nexamp)):
        totacc=[]
        for k in range(nfolds):
            dt=DecisionTree(purity[p],nexamp[n],7)#purity[p],nexamp[n])
            dt.train(folds[k][0][:,features],folds[k][1])
            pclasses=dt.predict(folds[k][2][:,features])
            acc=np.sum(pclasses==folds[k][3])/float(folds[k][3].shape[0])
            print ("[Info] Fold {} Accuracy = {}".format(k+1, acc))
            totacc.append(acc)
        params[p,n]=np.mean(totacc)
        print (totacc, '\nPurity={}, Nexample-threshold={}, Mean Accuracy ={}'.format(purity[p],nexamp[n], np.mean(totacc)))
```

Output:
```
Generating CV data for 2 classes
X.shape: (6093, 4)
X.shape: (2526, 4)
X.shape: (643, 4)
X.shape: (128, 4)
X.shape: (515, 4)
X.shape: (79, 4)
X.shape: (60, 4)
X.shape: (59, 4)
X.shape: (4, 4)
X.shape: (55, 4)
X.shape: (1, 4)
X.shape: (19, 4)
X.shape: (10, 4)
X.shape: (9, 4)
X.shape: (436, 4)
X.shape: (225, 4)
X.shape: (79, 4)
X.shape: (35, 4)
X.shape: (44, 4)
X.shape: (146, 4)
X.shape: (60, 4)
X.shape: (86, 4)
X.shape: (211, 4)
X.shape: (129, 4)
X.shape: (8, 4)
X.shape: (121, 4)
X.shape: (82, 4)
X.shape: (11, 4)
X.shape: (71, 4)
X.shape: (1883, 4)
X.shape: (1068, 4)
X.shape: (815, 4)
X.shape: (235, 4)
X.shape: (580, 4)
X.shape: (224, 4)
X.shape: (175, 4)
X.shape: (14, 4)
X.shape: (161, 4)
X.shape: (49, 4)
X.shape: (24, 4)
X.shape: (25, 4)
X.shape: (356, 4)
X.shape: (349, 4)
X.shape: (168, 4)
X.shape: (181, 4)
X.shape: (7, 4)
X.shape: (3567, 4)
X.shape: (1210, 4)
X.shape: (2357, 4)
X.shape: (629, 4)
X.shape: (597, 4)
X.shape: (579, 4)
X.shape: (352, 4)
X.shape: (231, 4)
X.shape: (121, 4
```

```python
np.unravel_index(np.argmax(params), params.shape)
```

Output:
```
(12, 0)
```